<a href="https://colab.research.google.com/github/glassresearch/PLT/blob/master/Python%20colab%20Georgia%20Tech/sort%20dict%20by%20value%20descend%20key%20ascend.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Problem 7

**Letter frequencies.** This problem has three (3) exercises worth a total of ten (10) points.

Letter frequency in text has been studied in cryptoanalysis, in particular frequency analysis. Linguists use letter frequency analysis as a rudimentary technique for language identification, where it's particularly effective as an indicator of whether an unknown writing system is alphabetic, syllablic, or ideographic.

Primarily, three different ways exist for letter frequency analysis. Each way generally results in very different charts for common letters. Based on the provided text, the first method is to count letter frequency in root words of a dictionary. The second way is to include all word variants when counting, such as gone, going and goes and not just the root word go. Such a system results in letters like "s" appearing much more frequently. The last variant is to count letters based on their frequency in the actual text that is being studied.

For more details, refer to the link:
https://en.wikipedia.org/wiki/Letter_frequency

In this problem, we will focus on the 3rd methodology.

**Exercise 0** (2 points). First, given a string input, define a function  `preprocess` that returns a string with non-alphabetic characters removed and all the alphabets converted into a lower case.

For example, 'We are coding letter Frequency! Yay!" would be transformed into "wearecodingletterfrequencyyay"

```python
def preprocess(S):
    ### YOUR CODE HERE

#   GOAL:
#   Return a string with non-alphabetic characters removed and all the alphabets converted into a lower case.

#   INPUT:
#   A string named 'S' to be processed.

#   STRATEGY:
#   1. Create an empty string. Let's call this 'final_string'
#   2. Get lowercase version of input 'S'. Let's call this 'S_lower'
#   2. Iterate over all of the characters in 'S_lower'
#        3. If the character is a letter, add it to 'final_string'
#   4. Return 'final_string'

#   My Google Search: 'python find if letter'
#   Link I found: https://www.geeksforgeeks.org/python-program-to-check-if-a-string-has-at-least-one-letter-and-one-number/

#   SOLUTION:
    final_string = ''
    S_lower = S.lower()
    '''
    for character in S_lower:
        if character.isalpha():
            final_string += character
    '''
    ret_string = "".join(letter for letter in S.lower() if letter.isalpha())
    return ret_string

    #return final_string

```

<!-- grade_id: valid_string, points: 2, grade: true, solution: false -->
```python
# Test cell: valid_string
import random, string

N_str = 100 #Length of random string

def generate_str(n):
    random_str = ''.join(random.choice(string.ascii_lowercase + string.ascii_uppercase + string.digits + string.punctuation) for _ in range(n))
    return random_str

def check_preprocess_str(n):
    random_str = generate_str(n)
    print("Input String: ",random_str)
    assert preprocess(random_str).islower() == True
    assert preprocess(random_str).isalpha() == True
    print("|----Your function seems to work correct for the string----|"+"\n")

check_preprocess_str(N_str)
check_preprocess_str(N_str)
check_preprocess_str(N_str)

print("\n(Passed)!")
```

Output:
```
Input String:  +,E@']uu8t|urw&R^"[<J}pezkKk6Kxospa6~0`ac*IT-DR{Ta&R*cj]^7L{.yD+FLf:p7Hj|c5Z1%aJ<@fK7[kil?s_q'""'i]%
|----Your function seems to work correct for the string----|

Input String:  %Yt*~Arf]$\SR"{dX@8bftXi}MX45Sckr9:PJF]~7gPr;Agaqr"^u#+u2[`IBccX=;a=5ddN.8}J_43vg>T0_$F/U~HNKCc-n;NO
|----Your function seems to work correct for the string----|

Input String:  Y"coh5PZCcyOE.b7;=RV7(lz[f=MBNf<lwqmRy.!%Yl?,]a^:6Tt<1#RS:Gb&-rz24}gzCjs\,f$Jg^0TuJ(JiX}dc<J(XOSI(m]
|----Your function seems to work correct for the string----|


(Passed)!
```

**Exercise 1** (4 points). With the necessary pre-processing complete, the next step is to write a function `count_letters(S)` to count the number of occurrences of each letter in the alphabet.  

You can assume that only letters will be present in the input string. It should output a dictionary and if any alphabet (a-z) is missing in the input string, it should still be a part of the output dictionary and its corresponding value should be equal to zero.


```python
def count_letters(S):
    ### YOUR CODE HERE

#   GOAL:
#   Count the number of occurrences of each letter in the alphabet.
#   Output a dictionary, and if any alphabet (a-z) is missing in the input string,
#   it should still be a part of the output dictionary,
#   and its corresponding value should be equal to zero.

#   INPUT:
#   A string 'S' which contains only letters.

#   STRATEGY:
#   1. Create an empty dictionary. Let's call this 'final_dict'
#   2. Create string with all letters a-z, which will become the keys in our final dictionary. Let's call this 'all_letters'
#   3. Iterate over each letter in 'all_letters' string.
#         4. Count the number of times that letter appears in our input string 'S'
#         5. Create a key, value pair in 'final_dict' where the key is the letter and its value is the count found in Step 4.
#   6. Return 'final_dict'

#   My Google Searches: 'python get all letters' and 'python count occurrences in string'
#   Links I found:
#   https://www.javatpoint.com/alphabet-in-python
#   https://www.programiz.com/python-programming/methods/string/count#google_vignette

#   SOLUTION:
    final_dict = {}

    # create the string of all lowercase letters
    all_letters = string.ascii_lowercase

    for letter in all_letters:
        final_dict[letter] = S.count(letter) # why don't we need to detect if the key is already there?\


    return final_dict
```

<!-- grade_id: count_letters, points: 4, grade: true, solution: false -->
```python
# Test cell: count_letters
import collections

N_processed_str = 100

def generate_processed_str(n):
    random_processed_str = ''.join(random.choice(string.ascii_lowercase) for _ in range(n))
    return random_processed_str

def check_count_letters(S):
    print("Input String: ",S)
    random_char = chr(random.randint(97,122))
    print("Character frequency evaluated for: ", random_char)
    if(random_char in S):
        assert count_letters(S)[random_char] == collections.Counter(S)[random_char]
        print("|----Your function seems to return correct freq for the char----|"+"\n")
    else:
        assert count_letters(S)[random_char] == 0
        print("|----Your function seems to return correct freq for the char----|"+"\n")

check_count_letters(generate_processed_str(N_processed_str))
check_count_letters(generate_processed_str(N_processed_str))
check_count_letters(generate_processed_str(N_processed_str))
print("\n(Passed)!")
```

Output:
```
Input String:  fcexwlfagwnhxdvejidykkfrckllixpqzmzuzuwvwequqqhppozrugesgpxyiuwgqsdwxehajwfksuivkhenrwargfortvgckizz
Character frequency evaluated for:  u
|----Your function seems to return correct freq for the char----|

Input String:  qrbptbqqourxammrqjugbbynevvrwgatijdgnlkdokijhkegtvxfeygcibsdfqijuravhedbtodmtoxazumdbmrqtiqqdvkrynbn
Character frequency evaluated for:  g
|----Your function seems to return correct freq for the char----|

Input String:  jolgukjrlfgqrdfszfecoghdobrzvyzgytyirddjiopdfpxatekncoojrrnwnlgsqmjlnewecezradgdtdhbzzrfcurjmvoxxphh
Character frequency evaluated for:  e
|----Your function seems to return correct freq for the char----|


(Passed)!
```

**Exercise 2** (4 points). The next step is to sort the distribution of a dictionary containing all the letters in the alphabet as keys and number of occurrences in text as associated value.

Sorting should be first done in decreasing order by occurrence count and for two elements with same count, the order should be alphabetic. The function  `find_top_letter(d)` should return the 1st character in the order.

```python
def find_top_letter(d):
    ### YOUR CODE HERE

#   GOAL:
#   Return the letter with the highest frequency in the input dictionary.
#   If two letters are tied in frequency, the one to return is the
#   one that appears first in the alphabet. So the letter to return will be the
#   key of the first element in the dictionary, after sorting.

#   INPUT:
#   A dictionary 'd' with letters as keys and frequencies as values

#   STRATEGY:
#   1. Figure out how to sort 'd' by values (frequency) in descending order and then by key (letter) in ascending order
#   2. Return the first letter (highest frequency, first in alphabetical order) from Step 1

#   My Google Searches: 'python sort dictionary' and 'python sort dictionary by value then key'
#   Links you could use:
#   A Fantastic Guide on Sorting Dictionaries in General:
#   https://realpython.com/sort-python-dictionary/

#   https://stackoverflow.com/questions/15371691/how-to-sort-a-dictionary-by-value-desc-then-by-key-asc
#   https://www.geeksforgeeks.org/python-sort-python-dictionaries-by-key-or-value/
#   https://www.geeksforgeeks.org/python-sort-dictionary-by-values-and-keys/
#   https://www.geeksforgeeks.org/python-get-the-first-key-in-dictionary/

#   SOLUTION:
    # First way:
    # kv[1] is the value → -kv[1] sorts values in descending order
    # - kv[0] is the key → sorts keys in ascending order if values are equal


    sorted_list = sorted(d.items(), key=lambda kv:(-kv[1], kv[0]))
    return sorted_list[0][0]

    # Second way:
    #sorted_dict = {val[0] : val[1] for val in sorted(d.items(), key = lambda x: (-x[1], x[0]))}
    #list_of_keys = list(sorted_dict.keys())
    #return list_of_keys[0]

```

<!-- grade_id: highest_freq_letter, points: 4, grade: true, solution: false -->
```python
# Test cell: highest_freq_letter

def create_random_dict():
    max_char_value = random.randint(5, 20)
    random_dict = {c:random.randint(0,max_char_value-1) for c in string.ascii_lowercase}
    random_letter1, random_letter2 = random.sample(string.ascii_lowercase, 2)
    random_dict[random_letter1], random_dict[random_letter2] = max_char_value, max_char_value
    if(random_letter1 < random_letter2):
        return random_letter1, random_dict
    else:
        return random_letter2, random_dict

def check_top_letter():
    top_letter, random_dict = create_random_dict()
    user_letter = find_top_letter(random_dict)
    assert user_letter == top_letter
    print("Input Dictionary: ", random_dict)
    print("Your function correctly returned most frequent letter: {} \n".format(user_letter))

check_top_letter()
check_top_letter()
check_top_letter()
print("\n(Passed)!")
```

Output:
```
Input Dictionary:  {'a': 10, 'b': 7, 'c': 3, 'd': 6, 'e': 5, 'f': 4, 'g': 5, 'h': 2, 'i': 3, 'j': 1, 'k': 10, 'l': 4, 'm': 11, 'n': 3, 'o': 8, 'p': 4, 'q': 8, 'r': 4, 's': 11, 't': 9, 'u': 2, 'v': 8, 'w': 3, 'x': 9, 'y': 5, 'z': 6}
Your function correctly returned most frequent letter: m 

Input Dictionary:  {'a': 12, 'b': 7, 'c': 13, 'd': 9, 'e': 10, 'f': 9, 'g': 13, 'h': 0, 'i': 8, 'j': 4, 'k': 6, 'l': 2, 'm': 6, 'n': 15, 'o': 9, 'p': 8, 'q': 0, 'r': 5, 's': 7, 't': 7, 'u': 1, 'v': 15, 'w': 7, 'x': 12, 'y': 2, 'z': 6}
Your function correctly returned most frequent letter: n 

Input Dictionary:  {'a': 2, 'b': 4, 'c': 3, 'd': 5, 'e': 3, 'f': 4, 'g': 3, 'h': 1, 'i': 3, 'j': 2, 'k': 0, 'l': 2, 'm': 4, 'n': 1, 'o': 1, 'p': 0, 'q': 1, 'r': 2, 's': 1, 't': 5, 'u': 4, 'v': 0, 'w': 3, 'x': 2, 'y': 0, 'z': 4}
Your function correctly returned most frequent letter: d 


(Passed)!
```

**Fin!** You've reached the end of this problem. Don't forget to restart the kernel and run the entire notebook from top-to-bottom to make sure you did everything correctly. If that is working, try submitting this problem. (Recall that you *must* submit and pass the autograder to get credit for your work!)